/*
TBD

// lifetimes3.rs
// Make this code compile.
// Execute `rustlings hint lifetimes3` for hints :)

// I AM NOT DONE

struct ABorrowingColor {
    // TODO: use a lifetime that is not 'static here
    name: &String,
    // TODO: use a lifetime that is not 'static here
    hex: &String, 
}

#[cfg(test)]
mod tests {
    use super::*;

    #[test]
    fn classic_c_structs() {
        let green = ColorClassicStruct { name: "green", hex: "#00FF00" };

        assert_eq!(green.name, "green");
        assert_eq!(green.hex, "#00FF00");
    }
}*/
